// variables4.rs
//
// Execute `rustlings hint variables4` or use the `hint` watch subcommand for a
// hint.

<<<<<<< HEAD
// ## mutable variable

fn main() {
    let mut x = 3;
=======
// I AM NOT DONE

fn main() {
    let x = 3;
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc
    println!("Number {}", x);
    x = 5; // don't change this line
    println!("Number {}", x);
}
