// variables5.rs
//
// Execute `rustlings hint variables5` or use the `hint` watch subcommand for a
// hint.

<<<<<<< HEAD
// ## variable shadowing
=======
// I AM NOT DONE
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc

fn main() {
    let number = "T-H-R-E-E"; // don't change this line
    println!("Spell a Number : {}", number);
<<<<<<< HEAD
    {
        let number = 3; // don't rename this variable
        println!("Number plus two is : {}", number + 2);
    }
=======
    number = 3; // don't rename this variable
    println!("Number plus two is : {}", number + 2);
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc
}
