// variables1.rs
//
// Make me compile!
//
// Execute `rustlings hint variables1` or use the `hint` watch subcommand for a
// hint.

<<<<<<< HEAD
// ## variable binding

fn main() {
    let x = 5u32;
=======
// I AM NOT DONE

fn main() {
    x = 5;
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc
    println!("x has the value {}", x);
}
