// variables3.rs
//
// Execute `rustlings hint variables3` or use the `hint` watch subcommand for a
// hint.

<<<<<<< HEAD
// ## variable declaration

fn main() {
    let x: i32;
    x = 10;
=======
// I AM NOT DONE

fn main() {
    let x: i32;
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc
    println!("Number {}", x);
}
