// variables6.rs
//
// Execute `rustlings hint variables6` or use the `hint` watch subcommand for a
// hint.

<<<<<<< HEAD
// # constant : 指定变量类型而不能省略

const NUMBER:i32 = 3;
=======
// I AM NOT DONE

const NUMBER = 3;
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc
fn main() {
    println!("Number {}", NUMBER);
}
