// variables2.rs
//
// Execute `rustlings hint variables2` or use the `hint` watch subcommand for a
// hint.

<<<<<<< HEAD

fn main() {
    let x = 10u32;
=======
// I AM NOT DONE

fn main() {
    let x;
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc
    if x == 10 {
        println!("x is ten!");
    } else {
        println!("x is not ten!");
    }
}
