// functions3.rs
//
// Execute `rustlings hint functions3` or use the `hint` watch subcommand for a
// hint.

<<<<<<< HEAD

fn main() {
    call_me(3);
=======
// I AM NOT DONE

fn main() {
    call_me();
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc
}

fn call_me(num: u32) {
    for i in 0..num {
        println!("Ring! Call number {}", i + 1);
    }
}
