// functions5.rs
//
// Execute `rustlings hint functions5` or use the `hint` watch subcommand for a
// hint.

<<<<<<< HEAD
// function bolck
=======
// I AM NOT DONE
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc

fn main() {
    let answer = square(3);
    println!("The square of 3 is {}", answer);
}

<<<<<<< HEAD
// solution 1
// fn square(num: i32) -> i32 {
//     return num * num;
// }

// solution 2 
fn square(num: i32) -> i32 {
    num * num
=======
fn square(num: i32) -> i32 {
    num * num;
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc
}
