// functions2.rs
//
// Execute `rustlings hint functions2` or use the `hint` watch subcommand for a
// hint.

<<<<<<< HEAD
// ## function parameter type 
// ## loop
=======
// I AM NOT DONE
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc

fn main() {
    call_me(3);
}

<<<<<<< HEAD
fn call_me(num:i32) {
=======
fn call_me(num:) {
>>>>>>> 24007d7b5f816a54e94386df6064a182a1cb66bc
    for i in 0..num {
        println!("Ring! Call number {}", i + 1);
    }
}
